# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

```python
import os
import pymysql
from sqlalchemy import create_engine
import pandas as pd
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!
# First, put in your user name below
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!

MYTABLE='dhumphrey420' + '_data'


# Then, create a connection to the same database, using the 
# same user name and password in the `mysql-examples.ipynb` 
# file.  Use conn as the database connection variable.

host = 'slucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
port = '3306'
user = 'slucor2020'
password = 'SLUcor2020'
database = 'hds5210'

conn = create_engine('mysql+pymysql://' + 
                     user + ':' + 
                     password + '@' + 
                     host + '/' + 
                     database, echo=False)

```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != 'paulboal_data')
assert(conn.connect())
```

```python
MYTABLE
```

Output:
```
'dhumphrey420_data'
```

```python

```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
url = 'http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv'

### SOLUTION GOES HERE
df= pd.read_csv(url)
df
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
```

```python
pd.read_sql_query('SHOW TABLES', conn)
```

```python
tables = {}

for n,t in pd.read_sql_query("SHOW TABLES", conn).iterrows():
    name = str(t['Tables_in_hds5210'])
    print(name)
    tables[name] = pd.read_sql("SELECT * FROM "+name,conn)
```

Output:
```
Medicaid_EP_Hospital_Type
NPI
Population
Provider_Name
Test
a85531647_data
aavila_data
aavila_data_NPI
aavila_data_new
aavila_data_year
aavila_datax
abbieslu
abbieslu_NPI
abbieslu_ipv
abbieslu_npi
abbieslu_year
achenliam
achenliam_data
achenliam_ipv
apuryear_data
arefehesmaeilpour_data
ashley_data
btabman_data
corona_counts
county_population
dhumphrey420_data
ed_visits
hospital_services
hospitals
jdgendron_data
jmccar23_data
kaylaschmidt_data
kenny0415_data
khalafh_data
kodurus_data
lchen_data_ipv
liannecota_data
liannecota_newdata
paulboal_data
petersluu_data
population
ppp_data
qichen_data
sam
sampele
sneha-koduru_data
/opt/tljh/user/lib/python3.6/site-packages/pymysql/protocol.py in check_error(self)
    218             errno = self.read_uint16()
    219             if DEBUG: print("errno =", errno)
--> 220             err.raise_mysql_exception(self._data)
    221 
    222     def dump(self):

/opt/tljh/user/lib/python3.6/site-packages/pymysql/err.py in raise_mysql_exception(data)
 
```

```python
data = tables['dhumphrey420_data'].merge(tables['population'])
```

```python
tables['dhumphrey420_data']
```

```python
dxyz
```

```python
dxyz.columns
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (328,19))
assert(list(dxyz.columns) == ['Provider_Name', 'NPI', 'CCN',
       'Medicaid_EP_Hospital_Type', 'Street_Address', 'City', 'County',
       'State', 'Zip_Code', 'Payment_Year_Number', 'Program_Type',
       'Total_payments', 'Last_Program_Year', 'Last_Payment_Year',
       'Last_Payment_Criteria', 'Most_Recent_Disbursement_Amount', 'LONGITUDE',
       'LATITUDE', 'FID'])
```

```python
join_table= pd.read_sql_query("""
  SELECT d.*, p.*
  FROM 
    dhumphrey420_data d JOIN
    population p ON d.Zip_Code = p.zipcode
    WHERE
    d.Zip_Code = p.zipcode
""", conn)
```

```python
join_table[['Zip_Code', 'zipcode']]
```

Output:
```
Zip_Code zipcode
0      90804   90804
1      91754   91754
2      91773   91773
3      91010   91010
4      91402   91402
..       ...     ...
82     91776   91776
83     91367   91367
84     90242   90242
85     91505   91505
86     90710   90710

[87 rows x 2 columns]
```

```python
match = join_table['Zip_Code'] == 'zipcode'
```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/pandas/core/ops/__init__.py:1115: FutureWarning: elementwise comparison failed; returning scalar instead, but in the future will perform elementwise comparison
  result = method(y)
```

```python
match.count
```

Output:
```
<bound method Series.count of 0     False
1     False
2     False
3     False
4     False
      ...  
82    False
83    False
84    False
85    False
86    False
Name: Zip_Code, Length: 87, dtype: bool>
```

```python
join_table['population'].max()
```

Output:
```
105549
```

```python
row=join_table.loc[join_table['population'] == 105549]
```

```python
row['NPI']
```

Output:
```
75    1194016923
Name: NPI, dtype: int64
```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Zip_Code` field and `population`'s `zipcode` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
join_table= pd.read_sql_query("""
  SELECT d.*, p.*
  FROM 
    dhumphrey420_data d JOIN
    population p ON d.Zip_Code = p.zipcode
    WHERE
    d.Zip_Code = p.zipcode
""", conn)
join_table['population'].max()
row=join_table.loc[join_table['population'] == 105549]
row['NPI']
answer = ['1194016923']
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)
assert(answer == ['1194016923'])
```

```python
dxyz.columns
```

Output:
```
Index(['Provider_Name', 'NPI', 'CCN', 'Medicaid_EP_Hospital_Type',
       'Street_Address', 'City', 'County', 'State', 'Zip_Code',
       'Payment_Year_Number', 'Program_Type', 'Total_payments',
       'Last_Program_Year', 'Last_Payment_Year', 'Last_Payment_Criteria',
       'Most_Recent_Disbursement_Amount', 'LONGITUDE', 'LATITUDE', 'FID'],
      dtype='object')
```

```python
join_table
```

Output:
```
Provider_Name         NPI    CCN  \
0                    LONG BEACH MEMORIAL MEDICAL  1720371065  50727   
1                AHMC GARFIELD MEDICAL CENTER LP  1730171265  50737   
2      PRIME HEALTHCARE SERVICES - SAN DIMAS LLC  1740447945  50588   
3           City of Hope National Medical Center  1750358297  50146   
4                          DEANCO HEALTHCARE LLC  1750365375  50704   
..                                           ...         ...    ...   
82    AHMC SAN GABRIEL VALLEY MEDICAL CENTER, LP  1275720377  50132   
83                   KAISER FOUNDATION HOSPITALS  1295880912  50677   
84      COUNTY OF LOS ANGELES AUDITOR CONTROLLER  1336154020  50717   
85  PROVIDENCE HEALTH SYSTEM-SOUTHERN CALIFORNIA  1336173269  50235   
86                   KAISER FOUNDATION HOSPITALS  1336294040  50411   

   Medicaid_EP_Hospital_Type       Street_Address            City  \
0       Acute Care Hospitals     1720 TERMINO AVE      LONG BEACH   
1       Acute Care Hospitals   525 N GARFIEL
```

```python
by_hospital= dxyz.groupby(['Zip_Code','Medicaid_EP_Hospital_Type'])['Total_payments'].sum().reset_index()
```

```python
by_hospital
```

Output:
```
Zip_Code Medicaid_EP_Hospital_Type  Total_payments
0       90004      Acute Care Hospitals         1292117
1       90012      Acute Care Hospitals         3840496
2       90015      Acute Care Hospitals         5580824
3       90017      Acute Care Hospitals         4195996
4       90018      Acute Care Hospitals         1553491
..        ...                       ...             ...
290     96104      Acute Care Hospitals          184615
291     96122      Acute Care Hospitals          515430
292     96130      Acute Care Hospitals         1323384
293     96150      Acute Care Hospitals          350678
294     96161      Acute Care Hospitals          527926

[295 rows x 3 columns]
```

```python
join_table= pd.read_sql_query("""
  SELECT b.*, p.*
  FROM 
     d JOIN
    population p ON d.Zip_Code = p.zipcode
    WHERE
    d.Zip_Code = p.zipcode
""", conn)
```

```python

```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
answer = pd.read_sql_query("""
Select Medicaid_EP_Hospital_Type, AVG(pay_per_capita) as Avg_Pay_per_Capita
From
(Select dhumphrey420_data.Medicaid_EP_Hospital_Type, dhumphrey420_data.Zip_Code, dhumphrey420_data.sum_tot_payments/population.population as pay_per_capita 
FROM
(SELECT Zip_Code, Medicaid_EP_Hospital_Type, SUM(Total_payments) as sum_tot_payments FROM
dhumphrey420_data
group by
Zip_Code, Medicaid_EP_Hospital_Type) dhumphrey420_data
JOIN
population population
ON
dhumphrey420_data.Zip_Code = population.zipcode) d
Group By
Medicaid_EP_Hospital_Type
""", conn)
answer
```

Output:
```
Medicaid_EP_Hospital_Type  Avg_Pay_per_Capita
0      Acute Care Hospitals        17974.362796
1      Children's Hospitals          111.559033
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 17974.363)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 111.559)

```

```python

```